---
title: Conditional Execution
abstract: |
  Conditional execution allows a program to selectively execute different parts of the code based on specific conditions. This capability enables the program to adapt and respond appropriately to various situations it encounters, making it more flexible and functional. By using boolean expressions and conditional statements, students will learn to direct the flow of execution to handle a wide range of inputs and scenarios effectively.
skip-execution: true
---

```python
import math

from flowcharts import *
from ipywidgets import interact
from dis import dis

%load_ext divewidgets
```

```python
if not input('Load JupyterAI? [Y/n]').lower()=='n':
    %reload_ext jupyter_ai
```

## Motivation

Conditional execution means running different pieces of code based on different conditions. Why do programmers need conditional executation?

For instance, when trying to compute `a/b`, `b` may be `0` and division by `0` is invalid.

```python
%%optlite -h 450
def multiply_or_divide(a, b):
    print("a:{}, b:{}, a*b:{}, a/b:{}".format(a, b, a * b, a / b))


multiply_or_divide(1, 2)
multiply_or_divide(1, 0)
```

Although division by 0 is invalid, multiplication remains valid but it is not printed due to the division error. Can we skip only the division but not the multiplication when `b` is equal to `0`?

**Solution 1**: Use a [*conditional expression*](https://docs.python.org/3/reference/expressions.html#conditional-expressions) that specifies which code block should be executed under what condition:

```python
... if ... else ...
```

```python
def multiply_or_divide(a, b):
    q = a / b if b else "undefined"
    print("a:{}, b:{}, a*b:{}, a/b:{}".format(a, b, a * b, q))


multiply_or_divide(1, 2)
multiply_or_divide(1, 0)  # multiplication is valid but not shown
```

**Solution 2**: Use a [*boolean expression*](https://docs.python.org/3/reference/expressions.html#boolean-operations):

```python
... and ... or ...
```

```python
def multiply_or_divide(a, b):
    q = b and a / b or "undefined"
    print("a:{}, b:{}, a*b:{}, a/b:{}".format(a, b, a * b, q))


multiply_or_divide(1, 2)
multiply_or_divide(1, 0)  # multiplication is valid but not shown
```

**Solution 3**: Monitor and catch the error using a [*try statement*](https://docs.python.org/3/reference/compound_stmts.html#try):

```python
def multiply_or_divide(a, b):
    try:
        q = a / b
    except ZeroDivisionError:
        q = "undefined"
    print("a:{}, b:{}, a*b:{}, a/b:{}".format(a, b, a * b, q))


multiply_or_divide(1, 2)
multiply_or_divide(1, 0)  # multiplication is valid but not shown
```

In this notebook, we will introduce the first two solutions. The last one is a better way to handle exceptions when the operations get complicated.

```python
%%ai
Explain in one paragraph, using a simple example, the benefits of using a 
try statement over conditional checks for handling exceptions and errors.
```

## Comparison Operators

A [*comparison/relational operators*](https://docs.python.org/3/reference/expressions.html#comparisons) along with its operands form an expression, which evaluates to a [boolean value](https://en.wikipedia.org/wiki/Boolean_data_type):

```python
True  # if the operands satisfy certain conditions, and
```

```python
False  # otherwise.
```

Unlike many other languages, Python capitalized the keywords for the boolean values to signify that they are constants, just like the keyword `None`. Hence,

```python
true = False  # is invalid in many languages but not Python
False = true  # is valid in many languages but not Python
```

::::{seealso} Why are boolean values "boolean"?
:class: dropdown

Boolean expressions are named after [George Boole](https://en.wikipedia.org/wiki/George_Boole), whose work laid the foundation of digital circuitry. Below are the links to the ebooks, thanks to [Project Gutenberg](https://www.gutenberg.org/about/):

- [Boole, George. The mathematical analysis of logic. CreateSpace Independent Publishing Platform, 1847.](https://www.gutenberg.org/ebooks/36884)
- [Boole, George. An investigation of the laws of thought: on which are founded the mathematical theories of logic and probabilities. Vol. 2. Walton and Maberly, 1854.](https://www.gutenberg.org/ebooks/15114)

An import result derived from Boole's framework for digital circuits is [Boole's/Shannon expansion theorem](https://en.wikipedia.org/wiki/Boole%27s_expansion_theorem).

::::

```python
%%ai
Explain Entscheidungsproblem and its relation to George Boole.
```

### Equality and Inequalties

The equality and inequality relationships in mathematics are implemented using the following comparison operators:

::::{table} Comparison operators for equality and inequalities
:label: tbl:math-comparison

| Expression |  True iff  |
| :--------- | :--------- |
|   `x == y`[^eq] | $x=y$.     |
|    `x < y`[^neq] | $x<y$.     |
|   `x <= y` | $x\leq y$. |
|    `x > y` | $x>y$.     |
|   `x >= y` | $x\geq y$. |
|   `x != y` | $x\neq y$. |

::::

[^eq]: The equality operator `==` consists of *two equal signs*, different from the assignment operator `=`.
[^neq]: Different from C or javascript:  
    - We can write `1 != 2` as `not 1 == 2` but not `!(1 == 2)` because
    - `!` is not a logical operator. It is used to call a [system shell command](https://ipython.readthedocs.io/en/stable/interactive/tutorial.html?highlight=system%20call#system-shell-commands) in IPython.

You can explore these operators using the widget below:

```python
comparison_operators = ["==", "<", "<=", ">", ">=", "!="]


@interact(operand1="10", operator=comparison_operators, operand2="3")
def comparison(operand1, operator, operand2):
    expression = f"{operand1} {operator} {operand2}"
    value = eval(expression)
    print(
        f"""{'Expression:':>11} {expression}\n{'Value:':>11} {value}\n{'Type:':>11} {type(value)}"""
    )
```

```python
%%ai
Explain very briefly how the comparison operators == may be represented 
differently in other programming languages such as maxima, scheme, and bash.
```

**What is the precedence of comparison operators?**

All the comparison operators have the [same precedence](https://docs.python.org/3/reference/expressions.html?highlight=precedence#operator-precedence) lower than that of `+` and `-`.

```python
1 + 2 >= 3  # (1 + 2) >= 3
```

Similar to the assignment operations, comparison operators can be chained together but they are [*non-associative*](https://en.wikipedia.org/wiki/Operator_associativity#Non-associative_operators).

```python
2.0 == 2 > 1, (2.0 == 2) > 1, 2.0 == (2 > 1)
```

::::{exercise} associativity
:label: ex:associativity

Explain why the following boolean expressions have different values.

::::

```python
1 <= 2 < 3 != 4, (1 <= 2) < (3 != 4)
```

<!-- grade_id: chained_relational_operator, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

::::{exercise} comparison operators
:label: ex:comparison

The comparison operators can be applied to different data types, as illustrated below.  
Explain the meaning of the operators in each of the following expressions.

::::

```python
# Comparisons beyond numbers
@interact(
    expression=[
        "10 == 10.",
        '"A" == "A"',
        '"A" == "A "',
        '"A" != "a"',
        '"A" > "a"',
        '"aBcd" < "abd"',
        '"A" != 64',
        '"A" < 64',
    ]
)
def relational_expression(expression):
    print(eval(expression))
```

<!-- grade_id: relational_expression, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Comparing Floating Point Numbers

`==` can also be used to compare floating point numbers:

```python
x = 10
y = (x ** (1 / 3)) ** 3
x == y
```

::::{caution} Why False? Shouldn't $(x^{\frac13})^3=x$?
:class: dropdown

Floating point numbers have finite precisions and so the computation $x^{\frac13}$ may not be exact.

::::

In practice, a small error in the calculation can be tolerated. Instead of checking equality of floating point numbers, we can check closeness up to a specified tolerance:

```python
abs_tol = 1e-9
y - abs_tol <= x <= y + abs_tol
```

`abs_tol`, often denoted as $\delta_{\text{abs}}$, is a positive number called the *absolute tolerance*.

**Why call it absolute tolerance?**

Note that the test remains unchanged if we swap `x` and `y`:

```python
abs_tol = 1e-9
x - abs_tol <= y <= x + abs_tol
```

Using the absolute function `abs`, we can also rewrite the comparison as follows:

```python
abs_tol = 1e-9
abs(x - y) <= abs_tol
```

::::{prf:definition} Absolute tolerance

The equality of $x$ and $y$ within an absolute tolerance of $\delta_{\text{abs}}$ means

$$x=y\pm \delta_{\text{abs}},$$

or more explicitly,

$$\lvert x-y\rvert \leq \delta_{\text{abs}}.$$

In other words, the comparison does not tolerate an absolute difference between $x$ and $y$ larger than $\delta_{\text{abs}}$.

::::

**Is an absolute tolerance of `1e-9` good enough?**

The same absolute tolerance fails if we set `x = 1e10` instead of `10`.

```python
x = 1e10
y = (x ** (1 / 3)) ** 3

abs_tol = 1e-9
abs(x - y) <= abs_tol
```

::::{caution} Why absolute tolerance fail?
:class: dropdown

In floating-point representation, the [radix point](https://simple.wikipedia.org/wiki/Radix_point) *floats* across different scales according to the exponent. The precision error in mantissa is therefore relative to the scale of the number.

::::

We should instead check whether $x$ is within a certain percentage of $y$:

```python
rel_tol = 1e-9
y * (1 - rel_tol) <= x <= y * (1 + rel_tol)
```

Note that the above test is not symmetric between $x$ and $y$, i.e., $x$ is relatively close to $y$ does not necessarily mean $y$ is relatively close to $x$.

```python
x = 1
y = 2
rel_tol = 0.5
y * (1 - rel_tol) <= x <= y * (1 + rel_tol), x * (1 - rel_tol) <= y <= x * (1 + rel_tol)
```

To make the comparison symmetric, we can use the absolute difference as follows:

::::{prf:definition} Relative tolerance

The equality of $x$ and $y$ is within a *relative tolerance* of $\delta_{\text{rel}}$ if

$$x=y \pm \delta_{\text{rel}} \max\Set{\lvert x\rvert,\lvert y\rvert},$$

or equivalently,

$$\lvert x-y\rvert \leq  \delta_{\text{rel}} \max\Set{\lvert x\rvert,\lvert y\rvert}.$$ (eq:rel_tol)

::::

::::{note} Why use $\max\Set{\lvert x\rvert,\lvert y\rvert}$ but not $\min\Set{\lvert x\rvert,\lvert y\rvert}$?
:class: dropdown

This is to make the test more tolerant, i.e.,

$$
\begin{align}
x&=y (1 \pm \delta_{\text{rel}}) && \text{ or } & y&=x (1 \pm \delta_{\text{rel}})
\end{align}
$$

instead of requiring both.

::::

::::{exercise} relative tolerance

Write the boolean expression that implements the above inequality. You can use the function `max(a, b)` to find the maximum of `a` and `b`.

::::

```python
x = 1e10
y = (x ** (1 / 3)) ** 3
```

<!-- grade_id: rel_tol, points: 0, grade: false, solution: true -->
```python
rel_tol = 1e-9
# YOUR CODE HERE
raise NotImplementedError
```

What if $x$ or $y$ is very close to 0? For instance, can you change `rel_tol` so that your solution above returns `True` for the following values of `x` and `y`?

```python
x = 1e-15
y = 0
```

::::{caution} When does relative tolerance fail?
:class: dropdown

If $x$ or $y$ is very close to 0, the relative tolerance test becomes ineffective, and so absolute tolerance is a more meaningful test.
In particular, with $y=0$, [](#eq:rel_tol) becomes 

$$
\begin{align}
\lvert x-0\rvert&\leq \delta_{\text{rel}} \max\Set{\lvert x\rvert, 0}\\
\lvert x\rvert&\leq \delta_{\text{rel}} \lvert x\rvert\\
1 &\leq \delta_{\text{rel}} \quad\text{or}\quad x=0.
\end{align}
$$

In the nontrivial case $x\neq 0$, the test $\delta_{\text{rel}}>1$ is independent of the value of $x$, and therefore cannot check whether $x$ is close to $0$. In particular, since $\delta_{\text{rel}}<1$ normally, and the test would always fail unless $x=0$ exactly, which defeats the purpose of allowing a tolerance.

::::

For a more flexible comparison, one can impose both the absolute and relative tolerances:

::::{prf:definition} absolute and relative tolerances combined

The following check combine both absolute and relative tolerances in one test

$$ \lvert x - y\rvert \leq \max\{\delta_{\text{rel}} \max\{\lvert x\rvert, \lvert y\rvert\},\delta_{\text{abs}}\}$$

::::

This is implemented by the function [`isclose`](https://docs.python.org/3/library/math.html#math.isclose) from `math` module with the default
- *relative tolerance* $\delta_{\text{rel}}$ equal to `1e-9`, and
- absolute tolerance $\delta_{\text{abs}}$ equal to `0.0`.

```python
math.isclose?
```

::::{exercise} `isclose` implementation
:label: ex:isclose_imp

Read the docstring carefully to implement your own `isclose` below. You can use the function `max(a, b)` to find the maximum of `a` and `b`.

:::{hint}
:class: dropdown

How to properly handle the floating point numbers `-inf`, `inf` and `NaN`? Note that those values are precise!

:::

::::

<!-- grade_id: isclose, points: 0, grade: false, solution: true -->
```python
def isclose(x, y, *, rel_tol=1e-09, abs_tol=0.0):
    # YOUR CODE HERE
    raise NotImplementedError
```

<!-- grade_id: test-isclose, points: 0, grade: true, solution: false -->
```python
x = 10
y = (x ** (1 / 3)) ** 3
assert isclose(x, y)
x = 1e10
y = (x ** (1 / 3)) ** 3
assert isclose(x, y)
assert not isclose(1e-15, 0)
assert isclose(float('inf'), float('inf'))
assert not isclose(float('nan'), float('nan'))
```

## Conditional Constructs

To illustrate how Python can carry out conditional execution, we will consider writing a program that sorts values in *ascending* order.

### If-Then Construct

**How to sort two values?**

Given two values are stored as `x` and `y`, we want to 
- `print(x,y)` if `x <= y`, and
- `print(y,x)` if `y < x`.

Such a program flow is often represented by a flowchart like the following:

```python
sort_two_values_fc1
```

**How to read the flowchart?**

A flowchart uses arrows to connects a set of [annotated blocks][bb]. The rules were first specified by ANSI and later adopted in [ISO 5807][iso].

[bb]: https://en.wikipedia.org/wiki/Flowchart#Building_blocks
[iso]: https://webstore.ansi.org/Standards/ISO/ISO58071985

**Why use a program flowchart?**

A program flowchart is a powerful way of describing an algorithm quickly. Unlike a text-based programming language:
- The rules governing the program flow can be shown explicitly by arrows.
- The annotated graphical blocks can convey the meaning faster using visual clues.

**How to implements the flowchart in Python?**

It is often useful to delay detailed implementations until we have written an overall skeleton. To leave a block empty, Python uses the keyword [`pass`](https://docs.python.org/3/tutorial/controlflow.html#pass-statements).

```python
# write a code skeleton
def sort_two_values(x, y):
    pass
    # print the smaller value first followed by the larger one


sort_two_values(1, 0)
sort_two_values(1, 2)
```

Python provides the [`if` statement](https://docs.python.org/3/reference/compound_stmts.html#the-if-statement) to implement the [*control flow*](https://en.wikipedia.org/wiki/Control_flow) specified by the diamond boxes in the flowchart.

```python
def sort_two_values(x, y):
    if x <= y:
        pass
        # print x before y
    if y < x: pass  # print y before x


sort_two_values(1, 0)
sort_two_values(1, 2)
```

To complete the implementations specified by the parallelogram boxes in the flow chart, we fill in the bodies/suites of the `if` statements as follows:

```python
def sort_two_values(x, y):
    if x <= y:
        print(x, y)
    if y < x: print(y, x)


@interact(x="1", y="0")
def sort_two_values_app(x, y):
    print("Values in ascending order:")
    sort_two_values(eval(x), eval(y))
```

Test the program by filling in different values of `x` and `y` above.

::::{seealso} How to indent code in Python?
:class: dropdown

As Python uses indentation to indicate code blocks or *suites*: 
- `print(x, y)` (Line 3) is indented to the right of `if x <= y:` (Line 2) to indicate it is the body of the if statement.
- For convenience, `if y < x: print(y, x)` (Line 4) is a one-liner for an `if` statement that only has one line in its block.
- Both `if` statements (Line 2-4) are indented to the right of `def sort_two_values(x,y):` (Line 1) to indicate that they are part of the body of the function `sort_two_values`.

The [style guide](https://www.python.org/dev/peps/pep-0008/#indentation) recommends using 4 spaces for each indentation. In JupyterLab, you can simply type the `tab` key.

::::

We can visualize the execution as follows. Step through the execution to

- see which lines are skipped, and
- understand why they are skipped.

```python
%%optlite -h 450
def sort_two_values(x, y):
    if x <= y:
        print(x, y)
    if y < x: print(y, x)


sort_two_values(1, 0)
sort_two_values(1, 2)
```

### If-Then-Else Construct

Can the sorting algorithm be improved further? 

::::{hint}
:class: dropdown

`not (x <= y) and (y < x)` is a *tautology*, i.e., always true.

::::

Consider the following modified flowchart:

```python
sort_two_values_fc2
```

This can implemented by the `else` clause of the [`if` statement](https://docs.python.org/3/tutorial/controlflow.html#if-statements) as follows:

```python
%%optlite -h 450
def sort_two_values(x, y):
    if x <= y:
        print(x, y)
    else:
        print(y, x)


sort_two_values(1, 0)
sort_two_values(1, 2)
```

Can we shorten the code to one line? This is possible with the [*conditional expression*](https://docs.python.org/3/reference/expressions.html#conditional-expressions).

```python
def sort_two_values(x, y):
    print(("{0} {1}" if x <= y else "{1} {0}").format(x, y))


@interact(x="1", y="0")
def sort_two_values_app(x, y):
    print("Values in ascending order:")
    sort_two_values(eval(x), eval(y))
```

::::{exercise}
:label: ex:conditional-expression

Explain why the followings have syntax errors.

```python
1 if True
x = 1 if True else x = 0
```

::::

<!-- grade_id: conditional-expression, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Nested Conditionals

Now, consider a slight more challenging problem of sorting three values instead of two. A feasible algorithm is as follows:

```python
sort_three_values_fc
```

To implement flowchart, we can use *nested conditional constructs*, where one conditional statement is placed within another, allowing for multiple layers of condition checks:

```python
def sort_three_values(x, y, z):
    if x <= y <= z:
        print(x, y, z)
    else:
        if x <= z <= y:
            print(x, z, y)
        else:
            if y <= x <= z:
                print(y, x, z)
            else:
                if y <= z <= x:
                    print(y, z, x)
                else:
                    if z <= x <= y:
                        print(z, x, y)
                    else:
                        print(z, y, x)


def test_sort_three_values():
    sort_three_values(0, 1, 2)
    sort_three_values(0, 2, 1)
    sort_three_values(1, 0, 2)
    sort_three_values(1, 2, 0)
    sort_three_values(2, 0, 1)
    sort_three_values(2, 1, 0)


test_sort_three_values()
```

Imagine what would happen if we have to sort many values. The program will not only be long, but also fat due to the indentation. To avoid an excessively long line due to the indentation, Python provides the `elif` keyword that combines `else` and `if`.

```python
def sort_three_values(x, y, z):
    if x <= y <= z:
        print(x, y, z)
    elif x <= z <= y:
        print(x, z, y)
    elif y <= x <= z:
        print(y, x, z)
    elif y <= z <= x:
        print(y, z, x)
    elif z <= x <= y:
        print(z, x, y)
    else:
        print(z, y, x)


test_sort_three_values()
```

::::{exercise}

The above sorting algorithm is inefficient because some comparisons may be checked more than once. Improve the program by:

1. Eliminating redundunt checks so it can run fast.
2. Eliminate redundunt variables so it does not require too much memory.
3. Making it short and easy to read.

:::{hint}
:class: dropdown
Do not use chained comparison operators or compound boolean expressions. E.g., 

```{code} python
:linenos:
:emphasize-lines: 10, 11, 13
def sort_three_values(x, y, z):
    if x <= y <= z:
        print(x, y, z)
    elif x <= z <= y:
        print(x, z, y)
    elif y <= x <= z:
        print(y, x, z)
    elif y <= z <= x:
        print(y, z, x)
    elif x <= y: # z <= x is redundant, why?
        print(z, x, y)
    else:
        print(z, y, x)
```

:::

::::

<!-- grade_id: efficient-sort_three_values, points: 0, grade: false, solution: true -->
```python
def sort_three_values(x, y, z):
    # YOUR CODE HERE
    raise NotImplementedError


sort_three_values(10, 17, 14)
```

::::{seealso} Sorting algorithms
:class: dropdown

A *sorting algorithm* refers to the procedure of sorting any number of values in order. The performance of a sorting algorithm is measured primarily by its [time complexity](https://en.wikipedia.org/wiki/Time_complexity) (the number of comparisons and swaps it makes) and its [space complexity](https://en.wikipedia.org/wiki/Space_complexity#:~:text=The%20space%20complexity%20of%20an,of%20characteristics%20of%20the%20input.) (the amount of additional memory it requires). It would be even better if the program is short and easy to read, but that depends partly on the design of the programming language. Python implemented its default `sorted` function using the [Timsort algorithm](https://en.wikipedia.org/wiki/Timsort#:~:text=Timsort%20is%20a%20hybrid%2C%20stable,in%20the%20Python%20programming%20language.). See if you can understand how it works and how well it works.

::::

```python
%%ai
Explain in one paragraph what the minimum possible number of comparisons needed
to sort n numbers is and why it is not the same as the maximum number of 
inversions.
```

## Boolean Operations

Since chained comparisons are non-associative, it follows a different evaluation rule than arithmetic operators.

E.g., `1 <= 2 < 3 != 4` is equivalent to:

```python
1 <= 2 and 2 < 3 and 3 != 4
```

The above is called a *compound boolean expression*, which is formed using the *boolean/logical operator* `and`.

**Why use boolean operators?**

What if we want to check whether a number is either $< 0$ or $\geq 100$?

The following program checks whether a number is either $< 0$ or $\geq 100$:

```python
# Check if a number is outside a range.
@interact(x="15")
def check_out_of_range(x):
    x_ = float(x)
    is_out_of_range = x_ < 0 or x_ >= 100
    print("Out of range [0,100):", is_out_of_range)
```

::::{seealso} Can you rewrite the logical expression to use only the `and` instead of `or`?
:class: dropdown

Turns out this is mission impossible because `and` alone is not [functionally complete](https://en.wikipedia.org/wiki/Functional_completeness),  i.e., it is not enough to give all possible boolean functions. We can add `or` and `not` operators to make it functionally complete.

::::

```python
%%ai
Explain in a paragraph why the set of operators including only "and" is not 
functionally complete.
```

```python
%%ai
Explain in a paragraph why the set of operators "and", "or", and "not" is 
functionally complete.
```

The following table is called a *truth table*. It enumerates all possible input and output combinations for each boolean operator available in Python:

::::{table} Truth table for different logical operators
:label: tbl:truth

|   `x`   |   `y`   | `x and y` | `x or y` | `not x` |
| :-----: | :-----: | :-------: | :------: | :-----: |
| `True`  | `True`  |  `True`   |  `True`  | `False` |
| `True`  | `False` |  `False`  |  `True`  | `False` |
| `False` | `True`  |  `False`  |  `True`  | `True`  |
| `False` | `False` |  `False`  | `False`  | `True`  |
::::

::::{note} What are the precedence and associativity for the logical operators?

- All binary boolean operators are left associative.  
- [Precedence](https://docs.python.org/3/reference/expressions.html?highlight=precedence#operator-precedence): `comparison operators` > `not` > `and` > `or`

::::

::::{exercise} compound boolean expressions
:label: ex:compound

Explain what the values of the following compound boolean expressions are:
- Expression A: `True or False and True`
- Expression B: `True and False and True`
- Expression C: `True or True and False`
::::

<!-- grade_id: compound-boolean, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

A compound boolean expression actually uses a [short-circuit evaluation](https://docs.python.org/3/reference/expressions.html?highlight=precedence#boolean-operations).

To understand this, we will use the following function to evaluate a boolean expression verbosely.

```python
def verbose(id, expr):
    """Identify evaluated boolean expressions."""
    print(id, "evaluated:", expr)
    return expr
```

For instance:

```python
verbose("1st expression", True)
verbose("2nd expression", False)
```

:::{seealso} Visualization by Thonny IDE
:class: dropdown

You may also use the debugger of [Thonny](https://thonny.org/) to visualize the step-by-step evaluations of the operands:

- In JupyterLab, click `File`$\to$`New Launcher`$\to$`Desktop` to launch a remote desktop for your server.[^1]
- Start Thonny in the remote desktop by clicking `Applications`$\to$`Development`$\to$`Thonny`.[^2]
- In the Thonny editor, write some Python code such as `True or False and True` and then save it to a file, say `test.py`.[^3]
- Debug the code by clicking the Thonny menu item `Run`$\to$`Debug current script (nicer)`.[^4]
- Click the `step into` button to step through the executions.

:::

[^1]: Thonny is an IDE based on TKinter, which requires a proper desktop environment to run.
[^2]: You can also launch a terminal in the remote desktop and run the command `thonny`.
[^3]: You can copy and paste code into the remote desktop by clicking `Remote Clipboard` button at the top right of the browser page and pasting your code into the textbox that appears below.
[^4]: Short-cut keys for Thonny may not work as the keys may be captured by the browser or your OS before they can reach Thonny.

### Short-circuit `or`

Consider evaluating `True or False and True` in a verbose manner:

```python
verbose("A", verbose(1, True) or verbose(2, False) and verbose(3, True))
```

::::{note}

When evaluating `True or False and True`,

- even though `or` has lower precedence than `and`, the short-circuit evaluation for `or` is performed first.
- The `and` operation is skipped entirely.

::::

**Why the second and third expressions are not evaluated?**

Because `True or ...` must be `True` (why?) so Python does not look further. From the [documentation](https://docs.python.org/3/reference/expressions.html?highlight=precedence#boolean-operations):

::::{card}
:header: Short-circuit evaluation of `or`

The expression `x or y` 
1. first evaluates `x`;  
2. if `x` is true, its value is returned;  
3. otherwise, `y` is evaluated and the resulting value is returned.

::::

Put it another way, `(x or y)` translate to the following

:::{code-block} python
:name: code:or
:caption: Translation of `(x or y)`

_ if (_ := x) else y

:::

except for the extra variable `_`. For example, `True or False and True` translates to `_ if (_ := True) else False and True`:

```python
verbose("A", (_ if (_ := verbose(1, True)) else verbose(2, False) and verbose(3, True)))
```

### Short-circuit `and`

Now, consider evaluating `False or False and True` in a verbose manner:

```python
False or False and True
```

```python
verbose("B", verbose(4, False) or verbose(5, False) and verbose(6, True))
```

**Why expression 6 is not evaluated?**

`False or False and ...` must be `False` so Python does not look further.

::::{card}
:header: Short-circuit evaluation of `and`

The expression `x and y` first evaluates `x`;  
if `x` is false, its value is returned;  
otherwise, `y` is evaluated and the resulting value is returned.
::::

Put it another way, `(x and y)` translate to the following

:::{code-block} python
:name: code:and
:caption: Translation of `(x and y)`

_ if not (_ := x) else y

:::


except for the extra variable `_`. For example, `False or False and True` translates to `False or (_ if not (_ := False) else True)`:

```python
verbose(
    "B", verbose(4, False) or (_ if not (_ := verbose(5, False)) else verbose(6, True))
)
```

### Non-Boolean Values for Boolean Operations and Control Flow Statements

Interestingly, Python also allows logical operations to have non-boolean operands and return values. From the [documentation](https://docs.python.org/3/reference/expressions.html?highlight=precedence#boolean-operations):

::::{card}
:header: Boolean interpretation

In the context of Boolean operations, and also when expressions are used by control flow statements, the following values are interpreted as false:  
- `False`
- `None`
- Numeric zero of all types
- Empty strings and containers (including strings, tuples, lists, dictionaries, sets and frozensets)

All other values are interpreted as true.

::::

The following is an application of this rule.

::::{exercise} short-circuit evaluation
:label: ex:short-circuit

How does the following code work?

::::

```python
print("You have entered", input() or "nothing")
```

<!-- grade_id: or, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

::::{exercise} chain
:label: ex:chain

Let's play a game to test your understanding of boolean expressions. Run the following program and choose your input so that 1, 2, 3 are each printed in a separate line as follows:
```bash
1
2
3
```

::::

```python
input(1) or input(2) and input(3)
```

<!-- grade_id: short_circuit_eval_game, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

::::{exercise} Chain using conditional expressions
:label: ex:chain-conditional-expr

Re-implement the following without any logical operators.

:::{code-block} python
:name: code:chain
:caption: Chaining inputs

input(1) or input(2) and input(3)

:::

:::{hint}
:class: dropdown

See [](#code:or) and [](#code:and).

:::

::::

<!-- grade_id: short_circuit_eval_to_conditional_exp, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError
```

::::{caution} Is an empty string `""` equal to `False`?
:class: dropdown

Try `"" == False`.

- An empty string is regarded as false in a boolean operation but
- a *comparison operation is not a boolean operation*, even though it forms a boolean expression.

Should the following return `False`?

```python
0 == False and 1 == True
```

Try and explain what you see.

::::

```python
%%ai
In Python, explain in one paragraph what is the value of the following expression:
0 == False and 1 == True
```