---
title: Functions
abstract: |
    Functions allow programmers to reuse code that is efficiently implemented and well-tested. This notebook not only demonstrates the basic syntax for using and writing functions, but also uses concrete examples to illustrate the importance of code reuse. It highlights the flexibility of functions by showing how they can be applied with different arguments to solve similar problems and how they can be customized to suit specific applications.
skip-execution: true
---

```python
import math
import ROOT

%reload_ext divewidgets
```

```python
if not input('Load JupyterAI? [Y/n]').lower()=='n':
    %reload_ext jupyter_ai
```

## What is a Function?

A function is a *callable* object, e.g.:

```python
callable(callable), callable(1)
```

The function `callable` is callable in the sense that
- it can be *called/invoked* with some input *arguments/parameters* such as `1` enclosed by parentheses `()`, and then
- *returns* some value computed from the input arguments, such as the boolean value `False` to indicate that the input argument `1` is not callable.

A function can be defined using the [`def` keyword](https://docs.python.org/3/tutorial/controlflow.html#defining-functions).

E.g., a simple function that prints "Hello, World!" can be defined as follows:

```python
# Function definition
def say_hello():
    print("Hello, World!")
```

```python
# Function invocation
say_hello()
```

To make a function more powerful and solve different problems,
- use a [return statement](https://docs.python.org/3/reference/simple_stmts.html#the-return-statement) to return a value that
- depends on some input arguments.

```python
def increment(x):
    return x + 1


increment(3)
```

A function must have a return value. By default `None` is returned.

```python
print(f"The return value is {say_hello()}.", )
```

We can also have multiple input arguments.

```python
def length_of_hypotenuse(a, b):
    return (a ** 2 + b ** 2) ** 0.5

length_of_hypotenuse(1, 2), length_of_hypotenuse(3, 4)
```

The arguments are evaluated from left to right:

```python
print("1st input:", input(), "\n2nd input:", input())
```

::::{note}

In summary, a function in Python
- may have any number (even zero) of input arguments, evaluated from left to right,
- may have *side effects* such as printing something to standard output, and
- must have a return value, which can be `None`.

::::

Indeed, how arguments are passed into a function can be more complicated than you may think. To check if you have the correct understanding: 

::::{caution} Does the code below increment `x` and prints `4`?
:class: dropdown

- Step 3: The function `increment` is invoked with an argument `x`.
- Step 3-4: A local frame is created for variables local to `increment` during its execution.    
    - The *formal parameter* `x` in `def increment(x):` becomes a local variable and
    - it is assigned the value `3` of the *actual parameter* given by the global variable `x`.
- Step 5-6: The local (but not the global) variable `x` is incremented.
- Step 6-7: The function call completes and the local frame is removed.

::::

```python
%%optlite -l -h 400
def increment(x):
    x += 1


x = 3
increment(x)
print(x)  # 4?
```

::::{seealso} Can we increment a variable instead of returning its increment?
:class: dropdown

In C++, it is possible to pass an argument by reference, allowing the local variable to point to the same memory location as the variable being passed in. Unlike Python, a variable in is C++ not merely a name; it is a named container whose size is determined by its type.

::::

```python
%%cpp
void increment(auto &x) {
    x += 1;
}

auto x = 3;
increment(x);
x
```

::::{seealso}

The above cell uses [`%%cpp` cell magic](https://github.com/root-project/root/blob/8aa112f13174ee8254092c52afb4653bd34e7eab/bindings/jupyroot/python/JupyROOT/kernel/magics/cppmagic.py) from [`ROOT` module](https://github.com/root-project/root) to run C++ code.

::::

```python
%%ai
Explain briefly the differences in how Python, C, and Java pass arguments to 
functions. In particular, explain
1. call by value,
2. call by reference, and
3. call by object reference.
```

A fundamental property of functions in Python is that they are [*first-class citizens*](https://en.wikipedia.org/wiki/First-class_function), which means that a function can be
1. assigned to a variable,
2. passed as an input argument, and
3. returned by a function.

```python
%%ai
Are there programming languages that do not treat functions as first-class citizens? Why?
```

The following is a simple illustration using the [`def` statement](https://docs.python.org/3/reference/compound_stmts.html#def) to define an `i`dentity function that uses the [`return` statement](https://docs.python.org/3/reference/simple_stmts.html#return) to return the input argument.

```python
%%optlite -h 300
def i(x):
    return x
assert i(i) == i and i.__name__ == 'i'
```

A function can also be defined using the [`lambda` expression](https://docs.python.org/3/reference/expressions.html#lambda), which creates an anonymous function:[^colon]

[^colon]: Note that the colon in `lambda ...:` cannot be followed by line break because it expects an expression rather than a suite.

```python
%%optlite -h 300
assert (i := lambda x: x)(i) == i \
and i.__name__ == "<lambda>"
```

A non-trivial example is the following implementation of the boolean values as functions:

```python
%%optlite -h 450
def true(x, y): return x
def false(x, y): return y
def ifthenelse(b, x, y): return b(x, y)
assert ifthenelse(true, "A", "B") == "A"
assert ifthenelse(false, "A", "B") == "B"
```

:::::{seealso} Why the name `<lambda>` for anonymous function?
:class: dropdown

The expression $\lambda x.M$, known as *function abstraction*, was defined by [Alonzo Church](https://en.wikipedia.org/wiki/Alonzo_Church), the Ph.D. supervisor of Alan Turing. This notation is used to define a function that can take (return) a string or another function passed in as $x$ (returned as $M$ but with $x$ substituted by the value of $x$). What is surprising is that a machine capable of taking and applying all such functions is [Turing complete](https://en.wikipedia.org/wiki/Turing_completeness). For more information, see the [$\lambda$-calculus](https://en.wikipedia.org/wiki/Lambda_calculus) and watch the following video created with `manim` or another video [here](https://youtu.be/eis11j_iGMs?si=X29joUgynqV3AtAO).

::::{card}
:header: [Open in a new tab](https://www.youtube.com/embed/ViPNHMSUcog?si=RfgWdE2gjUJW_Wqc)
:::{iframe} https://www.youtube.com/embed/ViPNHMSUcog?si=RfgWdE2gjUJW_Wqc
:::
::::

:::::

```python
%%ai
How to create an infinite loop using only the lambda expression in Python?
```

Perhaps you may also be interested in the following:

```python
%%ai
Why Alonzo Church used lambda for lambda calculus?
```

## Code Reuse

Previously, we learned about iteration, where the same piece of code can run multiple times. Function abstraction take this even further: It allows the same piece of code to be executed with different parameters and at different locations. Code reuse is a good programming practice. If done properly, it makes the code readible and efficient.  We will explore these benefits using a concrete example below.

### Perfect Square

::::{prf:definition} perfect square
:label: def:perfect_square

An integer $n$ is called a [perfect square](https://en.wikipedia.org/wiki/Square_number) if it is the square of an integer.

::::

For instance, the first 10 perfect squares are:

```python
for i in range(10):
    print(i**2)
```

Instead of generating perfect squares, how about writing a function that checks if a number is a perfect square?

::::{exercise}
:label: ex:perfect_square

Complete the following function to check if an integer `n` is a perfect square ([](#def:perfect_square)).

::::

<!-- grade_id: is_perfect_square, points: 0, grade: false, solution: true -->
```python
def is_perfect_square(n):
    # YOUR CODE HERE
    raise NotImplementedError
```

<!-- grade_id: test-is_perfect_square, points: 0, grade: true, solution: false -->
```python
# test cases
assert is_perfect_square(10**2)
assert not is_perfect_square(10**2 + 1)
assert is_perfect_square(10**10)
assert not is_perfect_square(10**10 + 1)
assert is_perfect_square(10**100)
assert not is_perfect_square(10**100 + 1)
```

As another demonstration of code reuse, the following solution uses a for loop to implement [](#def:perfect_square) exactly.

```python
def is_perfect_square(n):
    # checks if n is the square of i for i in the range up to n (exclusive). 
    for i in range(n):
        if i**2 == n:
            return True
```

If you try running the test on the above solution, it will take an unacceptably long time to run.[^interrupt] (Why?) 

[^interrupt]: Use the keyboard interrupt (&#9632;) to stop the execution.

To properly test the function, we should modify it to fail if it takes too long to run. Implementing such a feature, called *timeout*, is difficult. Fortunately, we can reuse the code written by others. Run the following cell to

1. install the package `wrapt_timeout_decorator` and
2. import the function `timeout` from the module `wrapt_timeout_decorator`.

```python
%pip install wrapt_timeout_decorator >/dev/null 2>&1
from wrapt_timeout_decorator import timeout
```

You will learn how to import a function in a subsequent section ([](#importing-external-modules)). For now, let's see how to use the `timeout` function:

```python
# enhanced test without timeout
duration = 5


@timeout(duration)  # raise error if the test does not complete in 5 seconds.
def test():
    if not input(f"Run the test with a timeout of {duration}s? [Y/n]").lower() == "n":
        assert is_perfect_square(10**2)
        assert not is_perfect_square(10**2 + 1)
        assert is_perfect_square(10**10)
        assert not is_perfect_square(10**10 + 1)
        assert is_perfect_square(10**100)
        assert not is_perfect_square(10**100 + 1)


test()  # run the test
```

::::{note} Does the for loop implementation pass all the test cases?
:class: dropdown

`assert is_perfect_square(10**10)` fails because Python is not fast enough to go through `10**10` numbers in `5` seconds.

::::

To add timeout to the test, we simply
1. wrapped the test inside a function `test`, and
2. decorated it with `@timeout(duration)`.

The function `test` will then be capabable of raising a `TimeOutError` if it takes more than the specified time `duration` to run.

You will learn how to write a [decorator](https://book.pythontips.com/en/latest/decorators.html#decorators) later. It is a powerful way to reuse functions and other objects with additional customizations.

### Integer Square Root

To improve the efficiency, consider the following sufficient and necessary condition for perfect squares:

::::{prf:proposition} integer square root
:label: pro:integer-sqrt

An integer $n$ is a perfect square iff

$$
n = (\lfloor \sqrt{n} \rfloor)^2,
$$ (eq:perfect_square)

namely, that the number is the square of its integer square root.

::::

A simple implementation is as follows:

```python
def is_perfect_square(n):
    # check if n is the square of its integer square root
    return n == int(n**0.5) ** 2

assert is_perfect_square(10**10)
```

Note that it fixed the efficiency issue on the test case with `n` being `10**10`. Let's run all the test cases:

```python
test()
```

::::{note} Does it pass all the test cases?
:class: dropdown

`assert is_perfect_square(10**100)` fails because of the finite precision of floating point numbers.

::::

Perhaps we should use `math.isclose` instead of `==`, since `int(n**0.5) ** 2` is a `float`.

```python
def is_perfect_square(n):
    return math.isclose(n, int(n**0.5) ** 2)


assert is_perfect_square(10**100)
```

Note that it can correctly say `10**100` is a perfect square. Let's run all the test cases:

```python
test()
```

::::{note} Does it pass all the test cases?
:class: dropdown

`assert not is_perfect_square(10**10+1)` fails because of the tolerance is too high.

::::

How to fix the issue? The culprit is that the computation for integer square root is not exact:

```python
x = 10**100
int((x) ** 0.5)
```

There are [better ways to compute integer square root](https://en.wikipedia.org/wiki/Integer_square_root). Binary search is a relatively easy one to try first, although it is not the best choice.

```python
%%ai
Explain very briefly how integer square root can be implemented in Python using
binary search.
```

But there is a much easier way to have a better implementation: Code reuse! Try `math.isqrt` for [](#ex:perfect_square) and check that you can pass all the test cases instantly.

```python
x = 10**100
math.isqrt(x), int((x) ** 0.5)
```

::::{seealso} How is `isqrt` implemented? 
:class: dropdown

`math.isqrt` implements an [adaptive-precision pure-integer version of Newton's iteration](https://github.com/python/cpython/blob/e5ab0b6aa68009a3f50b141ec013dacee3676db9/Modules/mathmodule.c#L1599C11-L1599C72).
The [source code](https://github.com/python/cpython/blob/e5ab0b6aa68009a3f50b141ec013dacee3676db9/Modules/mathmodule.c#L1791) is written in C, but there is a [Python implementation](https://github.com/python/cpython/blob/e5ab0b6aa68009a3f50b141ec013dacee3676db9/Modules/mathmodule.c#L1629) given in the source code, along with a [sketch of proof](https://github.com/python/cpython/blob/e5ab0b6aa68009a3f50b141ec013dacee3676db9/Modules/mathmodule.c#L1652).

::::

```python
%%ai
Explain briefly in two paragraph how isqrt is implemented as an 
adaptive-precision pure-integer version of Newton's iteration.
```

While you may want to write self-contained codes that do not rely on external libraries, code reuse advocates would recommend you to use standard libraries as much as possible. Why?

Indeed, the `math` library provides functions that it does not implement:

> **CPython implementation detail:** The `math` module consists mostly of thin *wrappers* around the platform C math library functions. - [pydoc last paragraph](https://docs.python.org/3/library/math.html)

E.g., see the [source code wrapper for `log`](https://github.com/python/cpython/blob/e5ab0b6aa68009a3f50b141ec013dacee3676db9/Modules/mathmodule.c#L757).[^CORDIC]:

[^CORDIC]: An efficient implementation often uses the [CORDIC algorithm](https://en.wikipedia.org/wiki/CORDIC).

```python
%%ai
When working on a programming assignment, should I write all the code myself,
or is it acceptable to use standard libraries?
```

## Modules

To facilitate code reuse, all Python codes are organized into libraries called *modules*. E.g., you can list all available modules using `pip list`:

```python
%pip list
```

Python searches for packages using the search path:

```python
import sys
sys.path
```

For instance, to show the location of a package, say `divewidgets`, run:

```python
%pip show divewidgets
```

::::{seealso}
:class: dropdown

See [](../Lab0/Setup.ipynb#Custom-Packages) for details of how to use `pip` and `conda` to install packages.
::::

### Builtins Module

In Python, every function must come from a module, including the build-in functions:

```python
__builtins__.print(f"`{print.__name__}` is from the {print.__module__} module.")
```

The `buildins` are automatically imported as `__builtins__` (and also `__builtin__`) along with all the functions and objects it provides because they are commonly use by programmers.

We can use the built-in function `dir` (*directory*) to list all built-in objects available.

```python
dir(__builtins__)
```

For instance, there is a built-in function `help` for showing the *docstring* (documentation string) of functions or other objects.

```python
help(help)  # can also show the docstring of help itself
```

```python
help(__builtins__)  # can also show the docstring of a module
```

::::{exercise}
:label: ex:dir 

We can call `dir` without arguments. What does it print?

:::{hint}
:class: dropdown

Try `help(dir)` or `dir?` in jupyter notebook.
:::

::::

```python
print(dir())
```

<!-- grade_id: dir, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

(importing-external-modules)=
### Importing External Modules

For other available modules, we can use the [`import` statement](https://docs.python.org/3/reference/simple_stmts.html#import) to import multiple functions or objects into the program *global frame*.

```python
%%optlite -h 300
from math import ceil, log10

x = 1234
print("Number of digits of x:", ceil(log10(x)))
```

The above imports both the functions `log10` and `ceil` from `math` to compute the number $\lceil \log_{10}(x)\rceil$ of digits of a *strictly positive* integer $x$.

Once can also import all functions from a library:

```python
%%optlite -h 300
from math import *  # import all except names starting with an underscore

print("{:.2f}, {:.2f}, {:.2f}".format(sin(pi / 6), cos(pi / 3), tan(pi / 4)))
```

The above uses the wildcard `*` to import ([nearly](https://docs.python.org/3/tutorial/modules.html#more-on-modules)) all the functions/variables provided in `math`.

::::{caution} What if different packages define the same function?
:class: dropdown

In the following code:

- The function `pow` imported from `math` overwrites the built-in function `pow`.  
- Unlike the built-in function, `pow` from `math` returns only floats but not integers or complex numbers. 
- We say that the import statement *polluted the namespace of the global frame* and caused a *name collision*.

::::

```python
%%optlite -h 500
print("{}".format(pow(-1, 2)))
print("{:.2f}".format(pow(-1, 1 / 2)))
from math import *

print("{}".format(pow(-1, 2)))
print("{:.2f}".format(pow(-1, 1 / 2)))
```

To avoid name collisions, it is a good practice to use the full name (*fully-qualified name*) such as `math.pow` prefixed with the module.

```python
%%optlite -h 350
import math

print("{:.2f}, {:.2f}".format(math.pow(-1, 2), pow(-1, 1 / 2)))
```

Using the full name can be problematic if the name of a module is very long. There can even be a hierarchical structure.
E.g., to plot a sequence using `pyplot` module from `matplotlib` package:

```python
%matplotlib widget
import matplotlib.pyplot

matplotlib.pyplot.stem([4, 3, 2, 1])
matplotlib.pyplot.ylabel(r"$x_n$")
matplotlib.pyplot.xlabel(r"$n$")
matplotlib.pyplot.title("A sequence of numbers")
matplotlib.pyplot.show()
```

In Python, modules can be structured into [packages](https://docs.python.org/3/tutorial/modules.html#packages), which are themselves modules that can be imported. It is common to rename `matplotlib.pyplot` as `plt`:

```python
import matplotlib.pyplot as plt

plt.stem([4, 3, 2, 1])
plt.ylabel(r"$x_n$")
plt.xlabel(r"$n$")
plt.title("A sequence of numbers")
plt.show()
```

We can also rename a function as we import it to avoid name collision:

```python
from math import pow as fpow

fpow(2, 2), pow(2, 2)
```

::::{exercise}
:label: ex:module_name 

What is wrong with the following code?

::::

```python
%%optlite -h 500
import math as m

for m in range(5):
    m.pow(m, 2)
```

<!-- grade_id: name-conflict, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Documentation

Understanding how to properly document a function is crucial for maintaining clear and efficient code. It also allow others to use the code properly to avoid bugs. How should one go about documenting a function effectively? As an example:

```python
# Author: John Doe
# Last modified: 2020-09-14
def increment(x):
    """Increment by 1.

    A simple demo of
    - parameter passing,
    - return statement, and
    - function documentation."""
    return x + 1  # + operation is used and may fail for 'str'
```

The `help` command shows the docstring we write 
- at the beginning of the function body
- delimited using triple single/double quotes.

```python
help(increment)
```

The docstring should contain the *usage guide*, i.e., information for new users to call the function properly. See Python style guide (PEP 257) for
- [one-line docstrings](https://www.python.org/dev/peps/pep-0257/#one-line-docstrings) and
- [multi-line docstrings](https://www.python.org/dev/peps/pep-0257/#multi-line-docstrings).

::::{note} Why doesn't `help` show the comments that start with `#`?
:class: dropdown

```python
# Author: John Doe
# Last modified: 2020-09-14
def increment(x):
    ...
    return x + 1  # + operation is used and may fail for 'str'
```

Those comments are not usage guide. They are intended for programmers who need to maintain/extend the function definition:

- Information about the author and modification date facilitate communications among programmers.
- Comments within the code help explain important and not-so-obvious implementation details.

::::

We can also [annotate](https://docs.python.org/3/library/typing.html) the function with *type hints* to indicate the types of the arguments and return value.

```python
# Author: John Doe
# Last modified: 2020-09-14
def increment(x: float) -> float:
    """Increment by 1.

    A simple demo of
    - parameter passing,
    - return statement, and
    - function documentation."""
    return x + 1  # + operation is used and may fail for 'str'


help(increment)
```

Annotations, if done right, can make the code easier to understand. However, annotations are not enforced by the Python interpreter.[^type-checking]

[^type-checking]: Type checking may be enforced by a supporting editor or packages such as [`pydantic`](https://docs.pydantic.dev/latest/) and [`mypy`](https://mypy.readthedocs.io/en/stable/).

```python
def increment_user_input():
    return increment(input())  # does not raise error even though input returns str
```

Does calling the function lead to any error:

```python
increment_user_input()
```

The types can also be described in the docstring following the [Numpy](https://numpydoc.readthedocs.io/en/latest/format.html#parameters) or [Google](https://google.github.io/styleguide/pyguide.html#383-functions-and-methods) style.

```python
# Author: John Doe
# Last modified: 2020-09-14
def increment(x: float) -> float:
    """Increment by 1.

    A simple demo of
    - parameter passing,
    - return statement, and
    - function documentation.

    Parameters
    ----------
    x: float
        Value to be incremented.

    Returns
    -------
    float:
        Value of x incremented by 1.
    """
    return x + 1  # + operation is used and may fail for 'str'


help(increment)
```

Can GenAI help us write documentations?

```python
%%ai
Add the document string to the following function:
--
def increment(x: float) -> float:
    return x + 1  # + operation is used and may fail for 'str'
```

::::{seealso} How to turn the docstrings into a user reference?
:class: dropdown

To faciliate the generation of documentation, there are tools such as:

- [sphinx](https://www.sphinx-doc.org/) that can compile the [docstrings automatically into an API reference](https://www.sphinx-doc.org/en/master/usage/extensions/autodoc.html).
- [`nbdev`](https://nbdev.fast.ai/) which releases a package and compiles the documentation from Jupyter notebooks, hence providing a literate programming experience.

::::