---
title: Symbolic Calculators
skip-execution: true
---

::::{attention}

This notebook is optional and NOT required for any course assessment activities. Lab tutor may go through them if time is available.

::::

```python
if not input('Load JupyterAI? [Y/n]').lower()=='n':
    %reload_ext jupyter_ai
```

Run the following to load the SymPy module for symbolic computations in Python.

```python
from sympy import *
```

The following code is a Python one-liner that creates a symbolic calculator.

```python
S(input())
```

Evaluate the cell repeatedly with `Ctrl+Enter` to try some calculations below using this calculator:

1. $2^3$ by entering `2**3`;
1. $\frac23$ by entering `2/3`;
1. $\left\lceil\frac32\right\rceil$ by entering `3//2`;
1. $3\mod 2$ by entering `3%2`;
1. $\sqrt{2}$ by entering `2**(1/2)`; and
1. $\sin(\pi/6)$ by entering `sin(pi/6)`;

Unlike `eval(input())`, JupyterLab render the $\LaTeX$[^latex] code of the expression using MathJax[^mathjax]. For instance, for the SymPy expression of $1/2$, the $\LaTeX$ code and the rendering by Mathjax is as follows:

[^latex]: [LaTeX](https://en.wikipedia.org/wiki/LaTeX) is a high-level typesetting system for creating structured documents with complex formatting for Mathematics as well. It is built on top of the [$\TeX$](https://en.wikipedia.org/wiki/TeX) typesetting system originally developed by Donald Knuth, with LaTeX itself created by Leslie Lamport.
[^mathjax]: [MathJax](https://en.wikipedia.org/wiki/MathJax) is a JavaScript library for displaying LaTeX, MathML, and AsciiMath notation in web browsers.

```python
sympy_expr = S("1/2")
print("LaTeX code:", sympy_expr._repr_latex_())
S("1/2")
```

For this lab, you will see how arbitrary precision arithmetic can be carried out using symbols.

## Symbolic Hypotenuse Calculator

We can define the following function to calculate the length `c` of the hypotenuse when given the lengths `a` and `b` of the other sides: 

:::{code-block} python
:name: code:symb:length_of_hypotenuse
:caption: A function that computes the length of hypotenuse
:linenos:
:emphasize-lines: 2

def length_of_hypotenuse(a, b):
    c = (a**2 + b**2) ** S("1/2")
    return c
:::

::::{caution}

`S("1/2")` is used instead of `1/2` for exact compution because `1/2` is a floating point value that has finite precision. Alternatively, you may replace
1. `S("1/2")` by `S(1)/2`, or
2. `( ... ) ** S("1/2")` by `sqrt( ... )`,

where `sqrt` is imported from the SymPy module.

::::

::::{exercise} length of hypotenuse

Complete the function below to return the exact length `c` of the hypotenuse given the lengths `a` and `b`.

:::{hint}
:class: dropdown

- See [](#code:symb:length_of_hypotenuse).

:::

::::

<!-- grade_id: length_of_hypotenus, points: 0, grade: false, solution: true -->
```python
def length_of_hypotenuse(a, b):
    # YOUR CODE HERE
    raise NotImplementedError
    return c
```

You can check your code against a few cases listed in the test cell below.

<!-- grade_id: test-length_of_hypotenus, points: 0, grade: true, solution: false -->
```python
# tests
assert length_of_hypotenuse(0, 0) == 0
assert length_of_hypotenuse(3, 4) == 5
assert length_of_hypotenuse(4, 7) == sqrt(65)
```

The tests use `==` instead of `isclose` because the compution should be exact. For instance, the hypotenuse of the last test case is:

```python
length_of_hypotenuse(4, 7)
```

<!-- grade_id: htest-length_of_hypotenus, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

If you are curious about the hidden test, it is like this:

:::{code-block} python
:name: code:htest-length_of_hypotenuse
:caption: Test of the exact formula of the length of hypotenuse
:linenos:
:emphasize-lines: 1

a, b = symbols('a, b', nonnegative=True)
assert length_of_hypotenuse(a, b) == sqrt(a**2 + b**2)
:::

The first line assign the variables `a` and `b` to the symbols `Symbol('a')` and `Symbol('b')` respectively, both of which are assumed to be nonnegative.

```python
a, b = symbols('a, b', nonnegative=True)
a, a.is_nonnegative, 0<=a<oo, sqrt(a).is_real, a<0, a>0
```

::::{note}

The variable `a` is assigned a new symbol with name `a` representing a (finite) nonnegative number, i.e., i.e., 

- $0\leq a<\infty$ holds, or equivalently
- $\sqrt{a} \in \mathbb{R}$.

Furthermore,
- `a<0` simplifies to `False` obviously as a is non-negative, but
- `a>0` does not simplify to `True` because `a` may or may not be `0`.

::::

Assumptions can affect the check for equality `==`:

```python
a_ = Symbol('a')
sqrt(abs(a_)**2) == a_, sqrt(abs(a) ** 2) == a
```

In other words,

- $\sqrt{\lvert a\rvert^2} \not\equiv a$ in general, but
- $\sqrt{\lvert a\rvert^2} = a \qquad \forall a\geq 0.$

But `==` does not check for mathematical equivalence:

```python
(a+1)*(a-1) == a**2-1, expand((a+1)*(a-1))==a**2-1
```

There is a further subtlety that, in Mathematics, equality `Eq` is different from `==`:

```python
Eq(sqrt(abs(a_)**2), a_), Eq(sqrt(abs(a) ** 2), a), Eq((a+1)*(a-1), a**2-1)
```

Further simplification is sometimes needed:

```python
simplify(Eq(sqrt(abs(a_)**2), a_)), simplify(Eq((a+1)*(a-1), a**2-1))
```

In the first case, the equality $\sqrt{\lvert a \rvert^2} = a$ 
- neither simplifies to `True`, as `a_` may be negative,
- nor simplifies to `False`, as `a_` may be non-negative.

In the second case, the equality holds `True` as `a` is non-negative.

Similarly, symbols with the same name but different assumptions may be equal but are not equivalent:

```python
simplify(Eq(a, a_)), a == a_, a == Symbol('a', nonnegative=True)
```

::::{seealso}
:class: dropdown

In SymPy, `==` and `Eq` are called the structural equality and symbolic equality respectively. For the implementation details, see the [SymPy documentation](https://docs.sympy.org/latest/explanation/glossary.html#term-Structural-Equality).

::::

## Calculus

Can we do complicated arithmetics with Python. What about Calculus?

$$
\int \tan(x)\, dx = \color{red}{?}
$$

Try [SymPy Beta](https://sympy-beta.vercel.app/input/integrate%28tan%28x%29%29):

- Take a look at the different panels to learn about the solution: `Steps`, `Plot`, and `Derivative`.
- Try different random examples.

::::{seealso} How does SymPy Beta work?
:class: dropdown

[SymPy Beta](https://github.com/eagleoflqj/sympy_beta) is a fork of [SymPy Gamma](https://github.com/sympy/sympy_gamma) that can run totally in the browser.

::::

While web applications like SymPy Beta offer a user-friendly and convenient interface for performing symbolic computations quickly without needing to install anything, using SymPy within a Python program provides significantly greater flexibility and control.

To compute the integration in Python, we first define a symbolic variable `x`:

```python
x = Symbol("x")
```

The SymPy expression for $\tan(x)$ is:

```python
f = tan(x)
f
```

To compute the integration:

$$
\int f(x) dx
$$

```python
c = Symbol("c")
g = integrate(f) + c  # c is assumed to be constant with respect to x
g
```

To compute the derivative:

$$
\frac{d}{dx}g(x)
$$

```python
diff_g = diff(g, x)
diff_g
```

The answer can be simplified as expected:

```python
simplify(diff_g)
```

```python
g.subs(c, 0)
```

To plot the functions:

```python
p = plot(f, g.subs(c, 0), (x, -2 * pi / 5, 2 * pi / 5), ylabel="y", legend=True)
```

::::{exercise} sympy
:label: ex:sympy

Using SymPy expressions

- assign to `x` a SymPy variable named `"x"`,
- assign to `f` the expression $\frac{1}{\sqrt{1 - x^2}}$ in terms of `x`,
- assign to `g` the result of $\int \frac{1}{\sqrt{1 - x^2}} dx$, and
- optionally, plot `f` and `g` for $x\in \left[-\frac45, \frac45\right]$.

:::{caution}
:class: dropdown

See [SymPy gotchas](https://docs.sympy.org/latest/gotchas.html).
:::

::::

<!-- grade_id: sympy, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError
```

The following test should plot your expression `f` in SymPy.

<!-- grade_id: test-sympy, points: 0, grade: true, solution: false -->
```python
# tests
assert simplify(f.subs(x, 0) - 1) == 0
assert simplify(f.subs(x, S(1)/2) - sqrt(3)*2/3) == 0
```

<!-- grade_id: htest-sympy, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

::::{seealso} How to program a symbolic calculator?

There is an excellent video lecture by Prof. Gerald Jay Sussman on [Symbolic Differentiation](https://ocw.mit.edu/courses/6-001-structure-and-interpretation-of-computer-programs-spring-2005/resources/3b-symbolic-differentiation-quotation/). The symbolic calculator he built can even do [variational calculus](https://en.wikipedia.org/wiki/Calculus_of_variations) for the computational approach to [Classical Mechanics](https://groups.csail.mit.edu/mac/users/gjs/6946/index.html) course. The textbook is available online:

> https://tgvaughan.github.io/sicm/

You can also play with the symbolic calculator in JupyterHub.

- Launch the Desktop app in JupyterHub: `File` → `New Launcher` → `Desktop`
- Open a terminal in the Desktop app: `Applications` → `Terminal Emulator`
- In the terminal, enter the following command, which will launch scheme along with the scmutils for symbolic computations:
    ```bash
    mechanics
    ```
- You can use an editor such as edwin/emacs to save your code to a file. Enter the following command into the scheme prompt to start Edwin:
    ```
    (edwin)
    ```
- Follow the [beginner's guide](https://groups.csail.mit.edu/mac/users/gjs/6.5160/beginner.pdf) or the [don't panic guide](https://groups.csail.mit.edu/mac/users/gjs/6946/dont-panic/#sec-6).



::::